# Manejo de archivos

### Ejemplo

Implementa una funci贸n llamada `multiplicacion` que reciba dos n煤meros ingresados por par谩metro y retorne la multiplicaci贸n de estos dos.

*Esta funci贸n deber铆a retornar un n煤mero*

```python
# Implementa tu respuesta en esta celda

def multiplicacion(num1, num2):
    
    respuesta = num1 * num2
    
    return respuesta
```

### Ejercicio 1

En la siguiente celda encuentras declarada una variable que almacena una cadena de caracteres.

```python
# No modifiques esta celda

cadena = "Archivos de texto"
```

Implementa una funci贸n llamada `sustituir` que reciba por par谩metro una cadena de texto (por ejemplo, la variable `cadena`) y retorne otra cadena en la cual se intercambien el segundo caracter y el pen煤ltimo caracter.

*Esta funci贸n debe retornar una cadena de texto.*

<!-- grade_id: ej1-sol, points: 0, grade: false, solution: true -->
```python
# your code here
def sustituir(cadena):
    if len(cadena) < 2:
        return cadena
    lista = list(cadena)
    lista[1], lista[-2] = lista[-2], lista[1]
    return ''.join(lista)
```

<!-- grade_id: ej1-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Variables base
cadena = "Archivos de texto"
cadena_prueba = "Una cadena con mayor cantidad de caracteres"

# Prueba 1: no existe la funci贸n
try:
    sustituir
    assert type(sustituir) == type(lambda:None)
except:
    raise NotImplementedError("No se encontr贸 una funci贸n llamada 'sustituir'.")

# Prueba 2: la funci贸n genera errores al ejecutarse
try:
    sustituir(cadena)
    sustituir(cadena_prueba)
except:
    raise RuntimeError("锔 Tu funci贸n lanza un error al ejecutarse. Revisa tu implementaci贸n.")

# Prueba 3: el resultado no es del tipo esperado
assert type(sustituir(cadena_prueba)) == str, f"Tu funci贸n debe retornar un valor de tipo {str.__name__}."

# Prueba 4: retorna un valor expl铆cito incorrecto
assert sustituir(cadena_prueba) != "Atchivos de texro", "Tu funci贸n est谩 devolviendo siempre la misma cadena fija. Debes usar el par谩metro."

# Prueba 5: la cadena no debe ser igual a la de entrada
assert sustituir(cadena_prueba) != cadena_prueba, "Parece que tu funci贸n no est谩 modificando la cadena recibida."

# Prueba 6: la longitud de la cadena no coincide
assert len(sustituir(cadena_prueba)) == len(cadena_prueba), "La cadena resultante no conserva la misma longitud."

# Prueba 7: segunda y pen煤ltima posici贸n no deben ser iguales
assert sustituir(cadena_prueba)[1] != sustituir(cadena_prueba)[-2], "Error: asignaste el mismo car谩cter en la segunda y pen煤ltima posici贸n."

# Prueba 8: error con la indexaci贸n de Python
assert sustituir(cadena_prueba)[2] != "s", "Revisa la indexaci贸n. En Python empieza desde 0."

# Prueba 9: comparaci贸n con el resultado esperado
assert sustituir(cadena_prueba) == "Uea cadena con mayor cantidad de caracterns", "La cadena devuelta no es la correcta."
assert sustituir(cadena) == "Atchivos de texro", "La cadena devuelta no es la correcta."

# Mensaje final
print(" 隆Excelente! Todas las pruebas fueron superadas correctamente.")
```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

### Ejercicio 2 

En la siguiente celda encuentras declarada una variable que almacena una cadena de caracteres representativa de la ruta relativa para el archivo `"tesoro.txt"`. Recuerda que para los talleres, los archivos necesarios est谩n en el subdirectorio `"Archivos"`.

```python
# No modifiques esta celda

directorio = "./Archivos/tesoro.txt"
```

Implementa una funci贸n llamada `archivo_a_lista` que reciba por par谩metro una cadena de texto que represente la direcci贸n de un archivo (por ejemplo, la variable `directorio`) y retorne una lista donde cada elemento es una linea del archivo encontrado.

*Esta funci贸n debe retornar una lista de cadenas de texto.*

<!-- grade_id: ej2-sol, points: 0, grade: false, solution: true -->
```python
# your code here
def archivo_a_lista(directorio):
    with open(directorio, encoding="utf-8") as archivo:
        return archivo.readlines()
```

<!-- grade_id: ej2-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Variables base
ruta = "./Archivos/tesoro.txt"
ruta_test = "./Archivos/tesoro_2.txt"

# Prueba 1: verificar que la funci贸n exista
try:
    archivo_a_lista
    assert callable(archivo_a_lista)
except:
    raise NotImplementedError("No se encontr贸 una funci贸n con el nombre 'archivo_a_lista'.")

# Prueba 2: ejecuci贸n sin errores
try:
    archivo_a_lista(ruta)
    archivo_a_lista(ruta_test)
except:
    raise RuntimeError("锔 La funci贸n genera un error al ejecutarse. Verifica que uses rutas relativas correctas.")

# Prueba 3: tipo de retorno
salida = archivo_a_lista(ruta_test)
assert isinstance(salida, list), "La funci贸n debe devolver una lista, no otro tipo de dato."

# Prueba 4: lista no vac铆a
assert len(salida) > 0, "La lista devuelta est谩 vac铆a. Aseg煤rate de leer correctamente el archivo."

# Prueba 5: validaci贸n de contenido inicial
primera_linea = archivo_a_lista(ruta_test)[0]
assert primera_linea != "Principales tenedores de t铆tulos del tesoro americano en lo que se lleva del 2020\n", "La respuesta es incorrecta para esta instancia."
assert archivo_a_lista(ruta)[0] != "Estimado estudiante.\n", "La respuesta es incorrecta. Debes utilizar el archivo recibido como par谩metro."

# Prueba 6: longitud de la lista
assert len(salida) == 8, "La cantidad de filas obtenidas no coincide con lo esperado."

# Prueba 7: contenido exacto
esperado = [
    "Estimado estudiante.\n",
    "\n",
    "Este archivo de texto es para uso exclusivo del auto-calificador.\n",
    "Evite borrar, mover o editar de cualquier forma su contenido.\n",
    "\n",
    "\n",
    "Agradecemos su colaboraci贸n,\n",
    "Equipo de Laboratorio Computacional de Analytics.\n"
]

assert salida == esperado, "El contenido de la lista no coincide con lo esperado."

# Mensaje final
print(" 隆Excelente! Todas las pruebas fueron superadas con 茅xito.")

```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

### Ejercicio 3 

En la siguiente celda encuentras declarada una variable que almacena una lista de paises y otra que guarda el resultado de ejecutar tu respuesta al [ejercicio 2](#Ejercicio-2).

```python
# No modifiques esta celda

lista_paises = ['Jap贸n', 'Francia']
lineas = archivo_a_lista(directorio)
```

Implementa una funci贸n llamada `filtrar_paises` que reciba por par谩metro una lista de nombres de paises (por ejemplo, la variable `lista_paises`) y retorne una lista donde cada elemento es la linea del archivo `"tesoro.txt"`, correspondiente a cada pa铆s de la lista que se recibe por par谩metro. Asegurate de haber realizado el ejercicio 2 correctamente y de usar la variable `lineas` en tu implementaci贸n para este ejercicio. A continuaci贸n, un ejemplo:

Ejecutar la funci贸n `filtrar_paises(["Suiza","Noruega"])` debe resultar en la siguiente lista:

```python
['Suiza\t247.4\t243.1\t241.3\t244.6\t243.7\t238.1\t237.5\t\n',
 'Noruega\t89.5\t87.6\t93.2\t98.0\t103.3\t97.7\t90.1\t\n']
```

*Esta funci贸n debe retornar una lista de cadenas de texto.*

<!-- grade_id: ej3-sol, points: 0, grade: false, solution: true -->
```python
# your code here
def filtrar_paises(lista):
    return [linea for linea in lineas if any(linea.startswith(pais + '\t') for pais in lista)]

```

<!-- grade_id: ej3-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
directorio = "./Archivos/tesoro.txt"
directorio_prueba = "./Archivos/tesoro_2.txt"
lista_paises = ['Jap贸n', 'Francia']
lista_paises_prueba = ['Suiza', 'Noruega', 'Tailandia', 'Kuwait']
resultado_prueba_ej_2 = ['Estimado estudiante.\n',
                         '\n',
                         'Este archivo de texto es para uso exclusivo del auto-calificador.\n',
                         'Evite borrar, mover o editar de cualquier forma su contenido.\n',
                         '\n',
                         '\n',
                         'Agradecemos su colaboraci贸n,\n',
                         'Equipo de Laboratorio Computacional de Analytics.\n']


# Prueba 0: no ha implementado el ejercicio 2:
try:
    assert archivo_a_lista(directorio_prueba) == resultado_prueba_ej_2
except:
    raise NotImplementedError("Debes haber implementado correctamente el ejercicio 2.")

lineas = archivo_a_lista(directorio)

# Prueba 1: no existe la funci贸n.
try:
    filtrar_paises
    assert type(filtrar_paises) == type(lambda:None)
except:
    raise NotImplementedError("No existe una funci贸n llamada filtrar_paises.",)
    
# Prueba 2: la funci贸n es interrumpida por errores durante su ejecuci贸n.
try:
    filtrar_paises(lista_paises)
    filtrar_paises(lista_paises_prueba)
except:
    raise RuntimeError("Tu funci贸n produce un error al ejecutarse.")

# Prueba 3: no retorna una lista
assert type(filtrar_paises(lista_paises_prueba)) == list, f"Tu funci贸n debe retornar un valor de tipo {list.__name__}."

# Prueba 4: retorna una lista vac铆a
assert len(filtrar_paises(lista_paises_prueba)) != 0, "Tu funci贸n retorna una lista vac铆a, cuando deber铆a devolver una lista con las filas solicitadas del archivo .txt."

# Prueba 5: respuesta explicita
assert filtrar_paises(lista_paises_prueba)[0] != 'Jap贸n\t1261.5\t1260.4\t1266.5\t1272.6\t1268.6\t1211.8\t1155.2\t\n', "Tu respuesta es incorrecta para una instancia diferente. Asegurate de utilizar el par谩metro."

# Prueba 6: devuelve una lista de longitud errada
assert len(filtrar_paises(lista_paises_prueba)) == len(lista_paises_prueba), "Tu funci贸n puede no estar tomando en cuenta todos los paises de la lista que recibe por par谩metro."

# Prueba 7: la lista es incorrecta
resultado = ['Jap贸n\t1261.5\t1260.4\t1266.5\t1272.6\t1268.6\t1211.8\t1155.2\t\n',
             'Francia\t144.2\t130.5\t137.7\t156.0\t147.6\t134.0\t127.7\t\n']

resultado_prueba = ['Suiza\t247.4\t243.1\t241.3\t244.6\t243.7\t238.1\t237.5\t\n',
                    'Noruega\t89.5\t87.6\t93.2\t98.0\t103.3\t97.7\t90.1\t\n',
                    'Tailandia\t85.5\t85.8\t80.5\t81.8\t91.3\t96.0\t90.5\t\n',
                    'Kuwait\t44.9\t43.6\t44.5\t40.1\t43.6\t41.9\t43.3\t\n']

assert filtrar_paises(lista_paises) == resultado, "Tu funci贸n devuelve una lista errada."
assert filtrar_paises(lista_paises_prueba) == resultado_prueba, "Tu funci贸n devuelve una lista errada."

# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")

```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

### Ejercicio 4

Jorge y Juan apostaron al que hiciera la mejor predicci贸n de la tabla final de las eliminatorias sudamericanas para el mundial de Rusia 2018. Jorge defini贸 el siguiente pron贸stico: 

|Puesto|<center>Pa铆s</center>|
|-------|-----------|
|1| Brasil|
|2| Argentina|
|3 | Uruguay|
|4| Colombia|
|5| Chile|
|6| Ecuador|
|7| Per煤|
|8| Paraguay|
|9| Venezuela|
|10| Bolivia|

Mientras que Juan defini贸 el siguiente pron贸stico:

|Puesto|<center>Pa铆s</center>|
|-------|-----------|
|1| Brasil|
|2| Uruguay|
|3 | Colombia|
|4| Argentina|
|5| Ecuador|
|6| Chile|
|7| Paraguay|
|8| Per煤|
|9| Venezuela|
|10| Bolivia|


El criterio elegido para decidir qui茅n gana la apuesta fue el error cuadr谩tico medio (ECM). El cual, para este ejemplo espec铆fico, tendr铆a la siguiente ecuaci贸n: 

$$
ECM= \frac{EC}{N} = \frac{\sum \limits _{j=1} ^{10} (Y_{j}-\hat{Y}_{j})^2}{10}
$$


Para calcular el ECM primero se calcula y eleva al cuadrado la diferencia entre la posici贸n predicha ($\hat{Y}_{j}$) y la posici贸n que realmente ocup贸 el equipo ($Y_{j}$). Por ejemplo, supongamos que esta fue la tabla final de las eliminatorias de Rusia 2018. Los errores cuadr谩ticos (EC) asociados con cada predicci贸n ser铆an los siguientes.


|Puesto|<center>Pa铆s</center>|EC(Jorge)|EC(Juan)|
|-----|--------|-----|-----|
|1|Brasil|0|0|
|2|Uruguay|1|0|
|3|Argentina|1|1|
|4|Paraguay|16|9|
|5|Colombia|1|4|
|6|Chile|1|0|
|7|Ecuador|1|4|
|8|Per煤|1|0|
|9|Venezuela|0|0|
|10|Bolivia|0|0|


Despu茅s se calcula el promedio de los errores cuadr谩ticos que en este caso ser铆a de 2.2 para la predicci贸n realizada por Jorge y de 1.8 para la predicci贸n realizada por Juan. Siendo que el error cuadr谩tico medio de Juan fue inferior al de Jorge, Juan ser铆a el ganador.

Jorge y Juan decidieron entonces usar la informaci貌n del art铆culo de Wikipedia (2020) sobre las eliminatorias a Rusia2018. Puntualmente copiaron y pegaron la "Tabla de Posiciones Final" disponible en el art铆culo, en un archivo de texto llamado `"Eliminatorias.txt"`. Este archivo contiene la tabla ordenada con la clasificaci贸n final de los pa铆ses as铆 como otros datos deportivos que en este caso no nos interesa analizar (puntos, goles a favor, goles en contra, etc.).

En la siguiente celda encuentras declaradas dos variables que almacenan en orden la clasificaci贸n pronosticada de Jorge y Juan respectivamente.

```python
# No modifiques esta celda

pronostico_jorge = ["Brasil", "Argentina", "Uruguay", "Colombia", "Chile", "Ecuador", "Per煤", "Paraguay", "Venezuela", "Bolivia"]
pronostico_juan = ["Brasil", "Uruguay", "Colombia", "Argentina", "Ecuador", "Chile", "Paraguay", "Per煤", "Venezuela", "Bolivia"]
```

Implementa una funci贸n llamada `ECM` que reciba por par谩metro una lista con el pr贸nostico de Juan, una con el pron贸stico de Jorge (en ese orden, por ejemplo, las variables `pronostico_jorge` y `pronostico_juan`). Tu funci贸n debe importar los resultados del archivo `"Eliminatorias.txt"` y retornar una tupla con el $ ECM $ de Jorge y Juan (en ese orden).

<!-- grade_id: ej4-sol, points: 0, grade: false, solution: true -->
```python
def ECM(pronostico_jorge, pronostico_juan):
    with open("./Archivos/Eliminatorias.txt", encoding="utf-8") as archivo:
        lineas = archivo.readlines()
    paises_reales = [linea.split('\t')[1].strip() for linea in lineas[1:] if linea.strip() != ""]
    n = len(paises_reales)
    ecm_jorge = sum((i - pronostico_jorge.index(pais)) ** 2 for i, pais in enumerate(paises_reales)) / n
    ecm_juan = sum((i - pronostico_juan.index(pais)) ** 2 for i, pais in enumerate(paises_reales)) / n
    return (ecm_jorge, ecm_juan)

```

<!-- grade_id: ej4-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
pronostico_jorge = ["Brasil", "Argentina", "Uruguay", "Colombia", "Chile", "Ecuador", "Per煤", "Paraguay", "Venezuela", "Bolivia"]
pronostico_juan = ["Brasil", "Uruguay", "Colombia", "Argentina", "Ecuador", "Chile", "Paraguay", "Per煤", "Venezuela", "Bolivia"]

pronostico_jorge_prueba = ["Argentina", "Chile", "Ecuador", "Per煤", "Uruguay", "Colombia", "Paraguay", "Venezuela", "Brasil", "Bolivia"]
pronostico_juan_prueba = ["Argentina", "Uruguay", "Brasil", "Ecuador", "Chile", "Paraguay", "Colombia", "Per煤", "Venezuela", "Bolivia"]

# Prueba 1: no existe la funci贸n.
try:
    ECM
    assert type(ECM) == type(lambda:None)
except:
    raise NotImplementedError("No existe una funci贸n llamada ECM.")
    
# Prueba 2: la funci贸n es interrumpida por errores durante su ejecuci贸n.
try:
    ECM(pronostico_jorge, pronostico_juan)
    ECM(pronostico_jorge_prueba, pronostico_juan_prueba)
except:
    raise RuntimeError("Tu funci贸n produce un error al ejecutarse.")

# Prueba 3: no retorna una tupla
assert type(ECM(pronostico_jorge_prueba, pronostico_juan_prueba)) == tuple, f"Tu funci贸n debe retornar un valor de tipo {tuple.__name__}."

# Prueba 4: retorna una tupla de tama帽o incorrecto
assert len(ECM(pronostico_jorge_prueba, pronostico_juan_prueba)) == 2, "Tu funci贸n debe retornar una tupla de dos elementos."

# Prueba 5: retorna una tupla pero no de floats
assert type(ECM(pronostico_jorge_prueba, pronostico_juan_prueba)[0]) == float and type(ECM(pronostico_jorge_prueba, pronostico_juan_prueba)[1]) == float, f"Los elementos de la tupla que retorna tu funci贸n no son de tipo {float.__name__}."

# Prueba 6: respuesta explicita
assert ECM(pronostico_jorge_prueba, pronostico_juan_prueba) != (1.4, 2.2), "La respuesta de tu funci贸n es incorrecta para una instancia diferente. Asegurate de utilizar el par谩metro."

# Prueba 7: orden incorrecto par谩metros
assert ECM(pronostico_juan, pronostico_jorge) == (2.2, 1.4), "La respuesta de tu funci贸n es correcta pero presenta los resultados en el orden equivocado."

# Prueba 8: valores incorrectos
assert ECM(pronostico_jorge, pronostico_juan) == (1.4, 2.2), "El c谩lculo de tus EMC es errado."
assert ECM(pronostico_jorge_prueba, pronostico_juan_prueba) == (12.8, 4.6), "El c谩lculo de tus EMC es errado."

# Mensaje de felicitaciones
print(" Felicidades, realizaste este ejercicio correctamente ")

```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

## CESAR ANDRES GARCES SUAREZ - 2220231028