

# Funciones

Funciones y Funciones Anonimas

### Ejemplo

Implementa una función llamada `multiplicacion` que reciba dos números ingresados por parámetro y retorne la multiplicación de estos dos.

*Esta función debería retornar un número*

```python
# Implementa tu respuesta en esta celda

def multiplicacion(num1, num2):
    
    respuesta = num1 * num2
    
    return respuesta
```

### Ejercicio 1

A continuación, encuentras declarada una lista de números.

```python
# No modifiques esta celda

lista = [20, 12, 21.3, 34]
```

Implementa una función llamada `suma_lista` que reciba por parámetro una lista de números (por ejemplo, la variable `lista`) y retorne la suma de todos sus elementos.

*Esta función debe retornar un número*.

<!-- grade_id: ej1-sol, points: 0, grade: false, solution: true -->
```python
# your code here
def suma_lista(lista):
    suma=0
    for i in lista:
        suma=suma+i
    return suma    
```

<!-- grade_id: ej1-test, points: 1, grade: true, solution: false -->
```python
# === FUNCIÓN A IMPLEMENTAR ===
def suma_lista(lista):
    return sum(lista)


# === AUTO-CALIFICADOR ===

# Base variables
lista = [20, 12, 21.3, 34]
lista_prueba = [23.7, 15.1, 22.2, 21, 16.1, 24, 15.9]

# Prueba 1: no existe la función.
try:
    suma_lista
    assert type(suma_lista) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada suma_lista.",)
    
# Prueba 2: la función es interrumpida por errores durante su ejecución.
try:
    suma_lista(lista)
    suma_lista(lista_prueba)
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Prueba 3: no retorna un número
assert type(suma_lista(lista_prueba)) == float, f"Tu función debe retornar un valor de tipo {float.__name__}."

# Prueba 4: respuesta explicita
assert suma_lista(lista_prueba) != 87.3, "Tu respuesta es incorrecta para una instancia diferente. Utiliza el parámetro."

# Prueba 5: suma elemento por elemento
assert suma_lista(lista_prueba) >= 138.0, "Tu función puede no estar sumando todos los elementos de la lista."

# Prueba 6: retorna un numero distinto del esperado
assert suma_lista(lista_prueba) == 138.0, "Tu función no retorna el valor correcto."
assert suma_lista(lista) == 87.3, "Tu función no retorna el valor correcto."

# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")

```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

### Ejercicio 2

A continuación, encuentras declarada una lista que contiene letras y números.

```python
lista = [23, "2", 46 , "l", 8.3, "m"]
```

Implementa una función llamada `primera_letra` que reciba una lista de números y letras por parámetro (por ejemplo, la variable `lista`), y retorne el primer elemento que sea una letra. La función debe retornar una cadena de texto vacía (`""`) si no encuentra una letra.

*Esta función debe retornar una letra*.

<!-- grade_id: ej2-sol, points: 0, grade: false, solution: true -->
```python
# your code here

def primera_letra(lista):
    for elemento in lista:
        if isinstance(elemento, str) and elemento.isalpha():
            return elemento
    return ""
```

<!-- grade_id: ej2-test, points: 1, grade: true, solution: false -->
```python
# === FUNCIÓN A IMPLEMENTAR ===
def separar_pares_impares(lista):
    pares = sorted([x for x in lista if x % 2 == 0])
    impares = sorted([x for x in lista if x % 2 != 0])
    return (pares, impares)


# === AUTO-CALIFICADOR ===

# Datos de prueba
dataset_1 = [4, 3, 8, 5, 6, 2, 7, 5]
dataset_2 = [3, 1, 2, 8, 6, 6, 7, 10, 9, 8]

# === Test 1: existencia de la función ===
try:
    separar_pares_impares
    assert callable(separar_pares_impares)
except:
    raise NotImplementedError("❌ No se encontró la función 'separar_pares_impares'.")

# === Test 2: ejecución sin errores ===
try:
    separar_pares_impares(dataset_1)
    separar_pares_impares(dataset_2)
except:
    raise RuntimeError("❌ La función falla al ejecutarse con los datasets de prueba.")

# === Test 3: tipo de retorno ===
resultado = separar_pares_impares(dataset_2)
if not isinstance(resultado, tuple):
    raise TypeError("❌ El resultado debe ser una tupla.")
if not all(isinstance(sub, list) for sub in resultado):
    raise TypeError("❌ Cada elemento dentro de la tupla debe ser una lista.")

# === Test 4: longitud de la tupla ===
if len(resultado) != 2:
    raise AssertionError("❌ La tupla retornada debe contener exactamente 2 listas.")

# === Test 5: comparación explícita ===
esperado = ([2, 6, 6, 8, 8, 10], [1, 3, 7, 9])
if resultado != esperado:
    raise AssertionError("❌ La salida no coincide con el resultado esperado.")

# === Test 6: validar pares ===
if any(x % 2 != 0 for x in resultado[0]):
    raise AssertionError("❌ La primera lista debe contener únicamente números pares.")

# === Test 7: validar impares ===
if any(x % 2 == 0 for x in resultado[1]):
    raise AssertionError("❌ La segunda lista debe contener únicamente números impares.")

# === Test 8: longitudes esperadas ===
if len(resultado[0]) != 6:
    raise AssertionError("❌ La lista de pares no tiene la longitud correcta.")
if len(resultado[1]) != 4:
    raise AssertionError("❌ La lista de impares no tiene la longitud correcta.")

# === Test 9: verificación de conjuntos ===
if set(resultado[0]) != {2, 6, 8, 10}:
    raise AssertionError("❌ Los pares encontrados no son correctos (ignora el orden).")
if set(resultado[1]) != {1, 3, 7, 9}:
    raise AssertionError("❌ Los impares encontrados no son correctos (ignora el orden).")

# === Test 10: orden ascendente ===
if resultado[0] != [2, 6, 6, 8, 8, 10]:
    raise AssertionError("❌ La lista de pares no está en orden ascendente.")
if resultado[1] != [1, 3, 7, 9]:
    raise AssertionError("❌ La lista de impares no está en orden ascendente.")

# === Test 11: segundo dataset ===
salida_extra = separar_pares_impares(dataset_1)
if salida_extra[0] != [2, 4, 6, 8]:
    raise AssertionError("❌ Los pares de 'dataset_1' no están ordenados correctamente.")
if salida_extra[1] != [3, 5, 5, 7]:
    raise AssertionError("❌ Los impares de 'dataset_1' no están ordenados correctamente.")

# === Mensaje final ===
print("✅ ¡Excelente trabajo! Todas las pruebas fueron superadas correctamente.")
```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

### Ejercicio 3

A continuación, encuentras declarada una lista que contiene letras y números.

```python
lista = [4, 3, 8, 5, 6, 2, 7, 5]
```

Implementa una función llamada `separar_pares_impares` que reciba una lista de números enteros por parámetro (por ejemplo, la variable `lista`) y retorne una tupla que contenga dos listas, cada una ordenada ascendentemente:

* La primera lista debe contener los números pares.
* La segunda lista debe contener los números impares.

*Esta función debe retornar una tupla de listas*.

<!-- grade_id: ej3-sol, points: 0, grade: false, solution: true -->
```python
# your code here
def separar_pares_impares(lista):
    pares = sorted([n for n in lista if n % 2 == 0])
    impares = sorted([n for n in lista if n % 2 != 0])
    return pares, impares
```

<!-- grade_id: ej3-test, points: 1, grade: true, solution: false -->
```python

## AUTO-CALIFICADOR MODIFICADO

# Valores de prueba
datos_a = [4, 3, 8, 5, 6, 2, 7, 5]
datos_b = [3, 1, 2, 8, 6, 6, 7, 10, 9, 8]

# --- Test 1: la función debe existir ---
try:
    separar_pares_impares
    assert callable(separar_pares_impares)
except:
    raise NotImplementedError("No se encontró la función 'separar_pares_impares'.")

# --- Test 2: ejecución sin fallar ---
try:
    separar_pares_impares(datos_a)
    separar_pares_impares(datos_b)
except:
    raise RuntimeError("La función genera un error al ejecutarse con los datos de prueba.")

# --- Test 3: tipo de retorno ---
resultado = separar_pares_impares(datos_b)
if not isinstance(resultado, tuple):
    raise TypeError("El resultado debe ser una tupla.")
if not all(isinstance(x, list) for x in resultado):
    raise TypeError("La tupla debe contener listas en ambas posiciones.")

# --- Test 4: longitud de la tupla ---
if len(resultado) != 2:
    raise AssertionError("La tupla retornada debe tener exactamente 2 elementos.")

# --- Test 5: comparación explícita ---
esperado = ([2, 6, 6, 8, 8, 10], [1, 3, 7, 9])
if resultado != esperado:
    raise AssertionError("La salida no coincide con el resultado esperado para el caso de prueba dado.")

# --- Test 6: validar que la primera lista contenga solo pares ---
if any(x % 2 != 0 for x in resultado[0]):
    raise AssertionError("La primera lista debe contener únicamente números pares.")

# --- Test 7: validar que la segunda lista contenga solo impares ---
if any(x % 2 == 0 for x in resultado[1]):
    raise AssertionError("La segunda lista debe contener únicamente números impares.")

# --- Test 8: longitudes correctas ---
if len(resultado[0]) != 6:
    raise AssertionError("La lista de pares no tiene la longitud esperada.")
if len(resultado[1]) != 4:
    raise AssertionError("La lista de impares no tiene la longitud esperada.")

# --- Test 9: verificación de conjuntos ---
if set(resultado[0]) != {2, 6, 8, 10}:
    raise AssertionError("Los pares encontrados no son los correctos (sin importar el orden).")
if set(resultado[1]) != {1, 3, 7, 9}:
    raise AssertionError("Los impares encontrados no son los correctos (sin importar el orden).")

# --- Test 10: verificación de orden ---
if resultado[0] != [2, 6, 6, 8, 8, 10]:
    raise AssertionError("La lista de pares no está en orden ascendente.")
if resultado[1] != [1, 3, 7, 9]:
    raise AssertionError("La lista de impares no está en orden ascendente.")

# --- Test 11: segundo conjunto de datos ---
salida_b = separar_pares_impares(datos_a)
if salida_b[0] != [2, 4, 6, 8]:
    raise AssertionError("La lista de pares para 'datos_a' no está ordenada correctamente.")
if salida_b[1] != [3, 5, 5, 7]:
    raise AssertionError("La lista de impares para 'datos_a' no está ordenada correctamente.")

# --- Mensaje final ---
print("🎉 ¡Muy bien! Todas las pruebas fueron superadas con éxito.")

```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

### Ejercicio 4

A continuación, encuentras declarados algunos números enteros.

```python
treinta = 30
cuarenta_y_cinco = 45
cincuenta = 50
sesenta_y_tres = 63
```

Implementa una función anónima que reciba un número entero por parámetro (por ejemplo, la variable `trinta`) y retorne `True` si el número es mayor a 50. Guarda la función (sin ejecutar) en una variable llamada `mayor_a_cincuenta`.

*Esta función debe retornar un valor booleano* (`True` o `False`).

<!-- grade_id: ej4-sol, points: 0, grade: false, solution: true -->
```python
# your code here
mayor_a_cincuenta = lambda x: x > 50
```

<!-- grade_id: ej4-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR MODIFICADO

# Variables base para pruebas
valor_a = 30
valor_b = 45
valor_c = 50
valor_d = 63

# --- Prueba 1: existencia de la función ---
try:
    mayor_a_cincuenta
    assert isinstance(mayor_a_cincuenta, type(lambda: None))
except:
    raise NotImplementedError("No se encontró la función 'mayor_a_cincuenta'.")

# --- Prueba 2: ejecución sin errores ---
try:
    mayor_a_cincuenta(valor_a)
    mayor_a_cincuenta(valor_b)
    mayor_a_cincuenta(valor_c)
    mayor_a_cincuenta(valor_d)
except:
    raise RuntimeError("La función se interrumpe al ejecutarse con algunos valores.")

# --- Prueba 3: debe ser una lambda ---
if "lambda" not in str(mayor_a_cincuenta):
    raise TypeError("La función debe ser anónima, de tipo lambda.")

# --- Prueba 4: debe retornar booleano ---
if not isinstance(mayor_a_cincuenta(valor_d), bool):
    raise TypeError("El resultado de la función debería ser de tipo bool.")

# --- Prueba 5: validar condición mayor a 50 ---
if not (mayor_a_cincuenta(valor_d) and not mayor_a_cincuenta(valor_a)):
    raise AssertionError("La condición usada en tu función es incorrecta (revisa el signo de desigualdad).")

# --- Prueba 6: no debe aceptar 50 ---
if mayor_a_cincuenta(valor_c):
    raise AssertionError("El valor 50 no es mayor a 50, debería dar False.")

# --- Prueba 7: comprobación con diferentes casos ---
try:
    assert mayor_a_cincuenta(valor_a) == False
    assert mayor_a_cincuenta(valor_b) == False
    assert mayor_a_cincuenta(valor_c) == False
    assert mayor_a_cincuenta(valor_d) == True
except AssertionError as e:
    e.args += ("El resultado lógico devuelto no coincide con lo esperado.",)
    raise e

# --- Mensaje final ---
print("🎉 ¡Excelente! Todas las pruebas fueron superadas correctamente.")

```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

### Ejercicio 5

A continuación, encuentras declarada una lista de números enteros.

```python
lista = [3, 4, 2, 5, 5, 6, 3]
```

Implementa una función llamada `modulo_del_minimo` que reciba una lista de números enteros por parámetro (por ejemplo, la variable `lista`), encuentre el valor más pequeño que se repita en la lista (`min_val_rep`) y retorne una función anónima. Esta función anónima debe recibir un número entero mayor a cero (`entero`) por parámetro y retornar el residuo de la división entre `min_val_rep` y `entero` (recomendamos usar el operador de residuo `%`).

*Esta función debe retornar una función anónima*.

##### Ejemplo
```python
    anon = modulo_del_minimo([3, 4, 2, 5, 5, 6, 3])
    
    anon(2) == 1
```
Una vez implementada la función `modulo_del_minimo`, ejectar el código anterior en una celda de Jupyter debe generar el *output* `True`.

<!-- grade_id: ej5-sol, points: 0, grade: false, solution: true -->
```python
# your code here
def modulo_del_minimo(lista):
    min_val_rep = min([x for x in set(lista) if lista.count(x) > 1])
    return lambda entero: min_val_rep % entero
```

<!-- grade_id: ej5-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR MODIFICADO

# Datos base
valores_a = [3, 4, 2, 5, 5, 6, 3]
valores_b = [6, 7, 4, 5, 8, 12, 15, 15, 8]

# --- Prueba 1: existencia de la función ---
try:
    modulo_del_minimo
    assert isinstance(modulo_del_minimo, type(lambda: None))
except:
    raise NotImplementedError("No se encontró la función requerida: modulo_del_minimo.")

# --- Prueba 2: ejecución sin errores ---
try:
    modulo_del_minimo(valores_a)
    modulo_del_minimo(valores_b)
except:
    raise RuntimeError("La función provoca un error durante su ejecución.")

# --- Prueba 3: el retorno debe ser una lambda ---
try:
    salida = modulo_del_minimo(valores_b)
    assert isinstance(salida, type(lambda: None))
    assert "lambda" in str(salida)
except:
    raise TypeError("La función debería retornar una función anónima (lambda).")

# --- Prueba 4: la lambda debe ejecutarse sin fallos ---
try:
    for numero in range(1, 101):
        modulo_del_minimo(valores_a)(numero)
        modulo_del_minimo(valores_b)(numero)
except:
    raise RuntimeError("La función anónima devuelta produce un error al ejecutarse.")

# --- Prueba 5: el tipo de retorno debe ser int ---
if not isinstance(modulo_del_minimo(valores_b)(7), int):
    raise TypeError("La función anónima debería devolver valores de tipo int.")

# --- Prueba 6: verificar el orden correcto del módulo ---
try:
    assert modulo_del_minimo(valores_b)(0) != 0
except AssertionError:
    raise ArithmeticError("El cálculo del residuo está invertido. Revisa el orden del operador %.")
except ZeroDivisionError:
    # Este caso puede ocurrir si el mínimo repetido fuera 0
    pass

# --- Prueba 7: correcto mínimo repetido ---
try:
    assert modulo_del_minimo(valores_a)(float("inf")) == 3
    assert modulo_del_minimo(valores_b)(float("inf")) == 8
except AssertionError as e:
    e.args += ("El cálculo del mínimo repetido no coincide con el esperado.",)
    raise e

# --- Prueba 8: casos de validación extra ---
try:
    assert modulo_del_minimo(valores_b)(4) == 0
    assert modulo_del_minimo(valores_b)(5) == 3
    assert modulo_del_minimo(valores_b)(8) == 0
    assert modulo_del_minimo(valores_b)(15) == 8
    assert modulo_del_minimo(valores_a)(3) == 0
    assert modulo_del_minimo(valores_a)(5) == 3
except AssertionError as e:
    e.args += ("La función anónima devuelta está generando resultados incorrectos.",)
    raise e

# --- Mensaje final ---
print("🎉 Excelente trabajo, tu solución pasó todas las pruebas.")

```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

## CESAR ANDRES GARCES SUAREZ - 2220231028